---
title: 'chapter 4'
description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 286cfb3c7f
xp: 50
```



`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 62c0351170
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```
