---
  title: "chapter 2"
  description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 4b86e00676   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 868bd164e9   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 5ff08884f7   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 5b4a6cb2de   
```

`@projector_key`
77d6aae5b6b0a63fa6b8f1678c4727fe
---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: ba4104edc1   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 67382445ba   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 2c5cdece95   
```





`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: b02c3e431f   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```

