---
  title: "Template Chapter 1"
  description: "This is a template chapter."
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: 850cee65b3   
```


Do some data science.


`@instructions`
instructions

`@hint`
hint

`@pre_exercise_code`

```{python}
# pre ex
```

`@sample_code`

```{sql}
SELECT * FROM `table`
```

`@solution`

```{sql}
SELECT * FROM `table`
```

`@sct`

```{python}
# sct
```

---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 461dd4a40e   
```

`@projector_key`
9709e80bbe7ff1f70eb817189c4d0c15