---
  title: "chapter 4"
  description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 286cfb3c7f   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 62c0351170   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```

