---
title: 'chapter 3'
description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: c286b02ff6
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 35625c0400
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 7913d7c751
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 4d18fa41c6
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: aac6253359
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: f406184bc0
xp: 50
```

`@projector_key`
440df5d798917f4729f7993b045536d2

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 18a5fc6e9f
xp: 50
```



`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 2d9af69a14
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 1d2fd1ab02
xp: 50
```



`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 47ffb5e519
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 3e512b6860
xp: 50
```



`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: b2d5da065c
xp: 50
```

`@projector_key`
c05373be38c00caba13a213f41b58591

---

## Insert exercise title here

```yaml
type: TabExercise
key: 3a553bcc15
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: c79e0d2626
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
