---
title: 'chapter 2'
description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 4b86e00676
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 868bd164e9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: TabExercise
key: 5ff08884f7
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 5b4a6cb2de
xp: 50
```

`@projector_key`
77d6aae5b6b0a63fa6b8f1678c4727fe

---

## Insert exercise title here

```yaml
type: TabExercise
key: ba4104edc1
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 67382445ba
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 2c5cdece95
xp: 50
```



`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: TabExercise
key: b02c3e431f
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```
