---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## d

```yaml
type: VideoExercise 
xp: 50 
key: 461dd4a40e   
```

`@projector_key`
9709e80bbe7ff1f70eb817189c4d0c15
---

## just changed

```yaml
type: NormalExercise 
xp: 100 
key: 4005c21732   
```


Goran was here!


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 8e4860d37a   
```





---

## New

```yaml
type: NormalExercise 
xp: 100 
key: 19850d30bc   
```





---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: f00e6e5fdd   
```





---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 3b1a7ad8af   
```





---

## Goran2

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 341acc154c   
```


this are some changes I want to do


---

## ta-da

```yaml
type: NormalExercise 
xp: 100 
key: 4f37d0fbe3   
```


goran test


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 860aca45d1   
```

`@projector_key`
40fe0e2b7f1a8672c99752eecb67dbae
---

## changed title

```yaml
type: NormalExercise 
xp: 100 
key: 7ff6daa0fc   
```


assignment


`@instructions`
instructions here

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: fe2f014aaf   
```





---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: ba7d3a19e7   
```





---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: d8d4c55f67   
```





---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 25841d5beb   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


