---
  title: "Template Chapter 1"
  description: "This is a template chapter."
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: 850cee65b3   
```


Do some data science.


`@instructions`
instructions

`@hint`
hint

`@pre_exercise_code`

```{python}
# pre ex
```

`@sample_code`

```{sql}
SELECT * FROM `table`
```

`@solution`

```{sql}
SELECT * FROM `table`
```

`@sct`

```{python}
# sct
```

---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 461dd4a40e   
```

`@projector_key`
9709e80bbe7ff1f70eb817189c4d0c15
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 4005c21732   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 8e4860d37a   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 3b1a7ad8af   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 19850d30bc   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 341acc154c   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: f00e6e5fdd   
```





`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 89ee7f102c   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 860aca45d1   
```

`@projector_key`
40fe0e2b7f1a8672c99752eecb67dbae
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 7ff6daa0fc   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: fe2f014aaf   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: ba7d3a19e7   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: d8d4c55f67   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```

