---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## d

```yaml
type: VideoExercise
key: 461dd4a40e
xp: 50
```

`@projector_key`
9709e80bbe7ff1f70eb817189c4d0c15

---

## just changed

```yaml
type: NormalExercise
key: 4005c21732
xp: 100
```

Goran was here!

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 8e4860d37a
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: d607f99548
xp: 50
```

`@projector_key`
d5966b7576e226cdda90f25375ef6bab

---

## New

```yaml
type: NormalExercise
key: 19850d30bc
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: f00e6e5fdd
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 3b1a7ad8af
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`
```{python}

```

---

## Goran2

```yaml
type: MultipleChoiceExercise
key: 341acc154c
xp: 50
```

this are some changes I want to do

`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 43b8341cda
xp: 100
```

dasljdbajsndklanlkdlaskdas

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 8cd1a3ca1d
xp: 50
```



`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


---

## ta-da

```yaml
type: NormalExercise
key: 4f37d0fbe3
xp: 100
```

goran test

`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: VideoExercise
key: 860aca45d1
xp: 50
```

`@projector_key`
40fe0e2b7f1a8672c99752eecb67dbae

---

## changed title

```yaml
type: NormalExercise
key: 7ff6daa0fc
xp: 100
```

assignment

`@instructions`
instructions here

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: fe2f014aaf
xp: 50
```



`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`

```{python}

```


`@possible_answers`


---

## Insert exercise title here

```yaml
type: TabExercise
key: ba7d3a19e7
xp: 100
```



`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d8d4c55f67
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: VideoExercise
key: 6a89a3c3b0
xp: 50
```

`@projector_key`
67c02f21220d638195c4717731d58a5c
